<a href="https://colab.research.google.com/github/hadassaholive/hds5210-2023/blob/main/week14/week14_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.  All of these files are found in our usual s3 bucket: `https://hds5210-data.s3.amazonaws.com`

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.

* https://hds5210-data.s3.amazonaws.com/npidata.csv


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

* https://hds5210-data.s3.amazonaws.com/cms2014.csv
* https://hds5210-data.s3.amazonaws.com/cms2015.csv
* https://hds5210-data.s3.amazonaws.com/cms2016.csv

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION

import pandas as pd
cms_2014 = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms2014.csv')
cms_2015 = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms2015.csv')
cms_2016 = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms2016.csv')
cms_2016.columns = map(str.lower, cms_2016.columns)
cms_2014['year'] = 2014
cms_2015['year'] = 2015
cms_2016['year'] = 2016
merged_cms = pd.concat([cms_2014, cms_2015, cms_2016], ignore_index=True)
if 'year' in merged_cms.columns:
    cols = [col for col in merged_cms.columns if col != 'year']
    reordered_cols = ['year'] + cols
    merged_df = merged_cms[reordered_cols]
else:
    print("The 'year' column is not found in the DataFrame.")
if 'Beneficiary_Average_Risk_Score' in merged_df.columns:
    merged_df.drop('Beneficiary_Average_Risk_Score', axis=1, inplace=True)
else:
    print("The 'Beneficiary_Average_Risk_Score' column is not found in the DataFrame.")
answer = merged_df

```

Output:
```
<ipython-input-1-40be0cd4a99b>:21: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  merged_df.drop('Beneficiary_Average_Risk_Score', axis=1, inplace=True)
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

```python

```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **https://hds5210-data.s3.amazonaws.com/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION

import pandas as pd
npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
answer = pd.merge(merged_df, npidata, left_on='npi', right_on='NPI', how='left')
```

Output:
```
<ipython-input-3-1de9a16ffb00>:6: DtypeWarning: Columns (18,33) have mixed types. Specify dtype option on import or set low_memory=False.
  npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
answer = None

if answer is not None:
    answer['State'] = answer['Provider Business Mailing Address State Name'].fillna('XX')

else:
    print("Please ensure 'answer' is assigned a DataFrame before executing this code.")








```

Output:
```
Please ensure 'answer' is assigned a DataFrame before executing this code.
```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-27-ead5e4353622> in <cell line: 1>()
----> 1 assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
      2 assert(answer.shape == (193862, 104))

AttributeError: 'NoneType' object has no attribute 'groupby'
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
answer = None

### SOLUTION
%matplotlib inline
answer = None
cms_table = cms.pivot_table(
    columns='year',
    index='State',
    values='total_services',
    aggfunc= 'sum'
)
answer = cms_table

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-39-78fe2b328a64> in <cell line: 7>()
      5 get_ipython().run_line_magic('matplotlib', 'inline')
      6 answer = None
----> 7 cms_table = cms.pivot_table(
      8     columns='year',
      9     index='State',

NameError: name 'cms' is not defined
```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-31-db10226393b0> in <cell line: 1>()
----> 1 assert(answer.shape == (4,3))
      2 assert(answer.sum().sum() == 519185664.6999999)
      3 assert(answer[2016].sum() == 176596933.80000004)
      4 assert(answer.loc['WY'].sum() == 10892707.4)

AttributeError: 'NoneType' object has no attribute 'shape'
```

```python
%matplotlib inline
answer.plot.bar()
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-32-5fc75f4fb76d> in <cell line: 2>()
      1 get_ipython().run_line_magic('matplotlib', 'inline')
----> 2 answer.plot.bar()

AttributeError: 'NoneType' object has no attribute 'plot'
```

```python

```